# Checkpoint 2
### Due: Tuesday, 7 November, 2023 at 11:00am GMT

### Total points: 100

### Read This First

Wherever you see raise NotImplementedError(), remove that line and put your code there.

Put the code that produces the output for a given task in the cell indicated. You are welcome to add as many cells as you like for imports, function definitions, variables, etc.

Your notebook must run correctly when executed once from start to finish. Your notebook will be graded based on how it runs, not how it looks when you submit it. To test this, go to the Kernel menu and select Restart & Run All.

Once you are happy with it, clear the output by selecting Restart & Clear Output from the Kernel menu.

### Your code must not take more than a minute per task.

Submit through Noteable.

# IMPORTANT 

In the following cell you must set the UID variable to be your student matriculation number (UUN), i.e.

`
UID=1769766
`

It must be an integer.

# Do not change it later

It will be used to generate the data that is different from other students



```python
UID =  # set it to your UUN number

# Checking it is integer
assert (type(UID) is int)
```

# Task 1 (20 points)

The first task is to find the number of local minima of the function. 

The function F() is defined in such a way that it accepts inputs of any number of dimensions. I.e. F([0]), F([0,1]), F([4,2,4,6,2]).

All the minima of the function are  located within -10..10 range. 
I.e. if the location of the minimum is at X=(x_1,x_2,...) then |x_i|<10 for every i.

Your task is to determine the number of local minima of the function for dimensions 1, 2, 3, 4, 5. 
Your code should return the list or array of 5 integer values.

You can assume that the function has no minima separated by less than 0.01.
I.e. if X is location of the local minimum and Y is a different local minimum then ||X-Y|| > 0.01

```python
import numpy as np
import scipy.optimize
import scipy.linalg
import matplotlib.pyplot as plt
import scipy.stats
%matplotlib inline
```

```python
T1_D = np.load('checkpoint2_task1.npz')['arr_0']
dims = [1, 2, 3, 4, 5]
import scipy.special as ssp


def blackbox_func(p):
    """
    This is is the function that you need to optimize
    It takes vector arguments (of any length) and a returns a scalar
    DO NOT CHANGE THE CODE INSIDE THIS FUNCTION
    """
    p = np.asarray(p)
    ndim = len(p)
    pos = dims.index(ndim)
    pos = np.random.default_rng(UID).permutation(5)[pos]
    curD = T1_D[1000 * pos:1000 * (pos + 1), :ndim]
    return -ssp.logsumexp(np.sum(-0.5 * ((p[None, :] - curD) / .3)**2, axis=1))


# The function is defined in  1,2,3,4,5 dimensions
print(blackbox_func([0]))
print(blackbox_func([0, 0]))
print(blackbox_func([0, 0, 0]))
```

<!-- grade_id: task1_solution, points: 0, grade: true, solution: true -->
```python
def find_number_of_minima():
    """
    Your function should return the list or array of 5 integers 
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: task1_test, points: 0, grade: true, solution: false -->
```python
my_minima = find_number_of_minima()
assert (len(my_minima) == 5)
print(my_minima)

```

# Task 2 (20 pts)

Generate random numbers from a probability distribution that looks like a triangle and a semi-circle.

The PDF f(x) is the following. Use the inverse CDF sampling method.

* f(x) = 0 when x<0
* f(x)= x/2 when 0<x<=1
* f(x)= 1-x/2 when 1<x<=2
* f(x)= 0 when 2<x<=3
* f(x)= $\pi/4 \sin((x-3)\cdot \pi)$ when 3<x<=4
* f(x)= 0 when x>4


<!-- grade_id: task2_solution, points: 0, grade: true, solution: true -->
```python
def generate_triangles(N):
    """ 
    Return a numpy array with N random numbers 
    following the distribution specified
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: task2_test, points: 0, grade: true, solution: false -->
```python
# The code will be tested by this
# the distribution of points should look like a triangle and semicircle
task2_result = generate_triangles(100000)
assert len(task2_result) == 100000
plt.hist(task2_result, bins=200)

```

# Task 3 (20 pts)

You are conducting an experiment where your parameter of interest x is subject to this equation 
$$ \int_0^x \left( a + sin(b  \ln(c + z^2) +e  z) \right) dz = y  $$


Your experiment provides you with the values of a, b, c, e, y and their uncertainties. 
You need to write the function that computes the value of x (5 points) and it's uncertainty (15 points)
You can assume that 0<x<100 and a>1 and c>0.

We will check the that the value of x is accurate to $0.0001$ and that uncertainty is within 10% of true uncertainty.

<!-- grade_id: task3_solution, points: 0, grade: true, solution: true -->
```python
def solve_task3(a=None,
                b=None,
                c=None,
                e=None,
                y=None,
                err_a=None,
                err_b=None,
                err_c=None,
                err_e=None,
                err_y=None):
    """
    This function will return two numbers x and its uncertainty given a, b, c, y values 
    and their uncertainties
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: task3_test, points: 0, grade: true, solution: false -->
```python
testme = solve_task3(
    a=3, b=1, c=0.2, e=0.4, y=5, err_a=0.2, err_b=0.2, err_c=0.02, err_e=0.1, err_y=0.2
)
print(testme)
assert len(testme) == 2
```

# Task 4 (20 pts)

# Fitting a periodic model to the data. 

In the input data file checkpoint2_task4_file.txt we have repeated measurements from the instrument. 
The file has 3 columns. 

One is the time of observation, the second one is the measurement, and the third one is the uncertainty of the measurement.  We think that the measurements should show periodic changes with time, so we want to fit the dataset by a model with a periodic component:
$$ M(t) = A \sin(2\pi t/T) + B \cos(2\pi t/T) + C t +D $$
A, B, C, D, T are model parameters. We are interested in determining T.

Assuming that we know that the period is between 0.5 and 5, determine the best period describing the data with the accuracy of better than 0.01.

Your function needs to return the pair of  
* the best period and 
* model values of your best fit model corresponding to the measurement times in the input file.

Hint: Notice the model is linear with respect to parameters  A,B,C,D

The grading will be done on another data file you do not have access to.

<!-- grade_id: task4_solution, points: 0, grade: true, solution: true -->
```python
def solve_period_model(filename):
    """ 
        This function should read the data from the file
        Fit the model and
        return the pair of best-fit period, best-fit model value array 
    """
    DATA = np.loadtxt(filename)
    t, y, ey = DATA
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: task4_test, points: 0, grade: true, solution: false -->
```python
# Your code will be tested like this on the provided data 
# AND on the hidden dataset
bestper, bestmod = solve_period_model('checkpoint2_task4_file.txt')
T4_DATA = np.loadtxt('checkpoint2_task4_file.txt')
plt.errorbar(T4_DATA[0], T4_DATA[1], T4_DATA[2], fmt='.')
print('BEST PERIOD', bestper)
plt.plot(np.sort(T4_DATA[0]), bestmod[np.argsort(T4_DATA[0])])

```

# Task 5 (20 pts)

You are given a dataset from the particle physics experiment. 
In that experiment the number of particles are recorded in intervals of energy.
The dataset contains an array of those numbers and the central energy of each bin/interval. 
We think that the average number of particles in each bin could be described by this model 

$$R(E) = |A_1|+ 100 \cdot \exp\left(A_2 \cdot \frac{E-70}{100}\right) + |A_3| \cdot  \exp
\left( - \frac{1}{2}\frac{(E-A_4)^2}{A_5^2} \right)$$

where the first two terms are describing the background and the last term is a contribution from a decaying 
particle with the mean energy $A_4$. 

In the dataset you can assume that the number of particles in each bin is described by Poisson distribution with the rate described by the equation above. You can assume that measurements of number of particles in each bin are independent from each other.

Your task is to model the data, estimate the energy of the particle $A_4$ using maximum likelihood approach and estimate the uncertainty of your measurement.

Your function solve_task5() should return two numbers: the estimate of the energy and its uncertainty. 
If you don't know how to estimate the uncertainty, you can return np.nan instead

Your function should also overplot your best model on top of the data.

Hint: Be careful with exponent overflow! (The value of $A_2$ should not be too large!)

We will be testing your code using the same data that you see.

```python
def task5_data_reader(filename):
    # This simply reads the data and return the array of Energies and the number counts in each bin
    # Use this function to read the data
    # DO NOT CHANGE THIS
    D = np.loadtxt('checkpoint2_task5.txt',
                   dtype=[('i', float), ('E', np.float64), ('N', float)])
    i = D['i'].astype(int)
    E, N = D['E'][i == UID % 100], D['N'][i == UID % 100].astype(int)
    return E, N

T5_E, T5_N = task5_data_reader('checkpoint2_task5.txt')
plt.plot(T5_E, T5_N, drawstyle='steps')
plt.xlabel('Energy')
plt.ylabel('Number of particles')
```

<!-- grade_id: task5_solution, points: 0, grade: true, solution: true -->
```python
def solve_task5(filename):
    """
    Your function should return the best value of A4 and its uncertainty. 
    It also should plot the data and overplot your best model on top of the data
    """
    T5_E, T5_N = task5_data_reader(filename)
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: task5_test, points: 0, grade: true, solution: false -->
```python
bestval, besterr = solve_task5('checkpoint2_task5.txt')

```